<a href="https://colab.research.google.com/github/Riab22209/CE880-Lab-Task-repository/blob/main/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

---
# Welcome to CE880
### This is your week-10 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 



```python
import numpy as np
from sklearn import datasets, linear_model
from sklearn.metrics import mean_squared_error, r2_score

# Load the diabetes dataset
diabetes_X, diabetes_y = datasets.load_diabetes(return_X_y=True)

# Use only one feature
diabetes_X = diabetes_X[:, np.newaxis, 2]

# Split the data into training/testing sets
diabetes_X_train = diabetes_X[:-20]
diabetes_X_test = diabetes_X[-20:]

# Split the targets into training/testing sets
diabetes_y_train = diabetes_y[:-20]
diabetes_y_test = diabetes_y[-20:]
```

---
## Question 1: 

#### Regression
Use the diabetes dataset from sklearn, which I have spilted into training and test sets (diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), please train the linear regression model to get the test accuracy and return `mean squared error`.  

<!-- grade_id: my_reg_model_mse_ans, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error
def my_reg_model_mse(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test):
    """Write a function to train a machine learning model to get the `mean_squared_error` on diabetes_X_test. 
    Please use LinearRegression model to train a regressor. """
    # YOUR CODE HERE
    pred_model = LinearRegression()
    pred_model.fit(diabetes_X_test, diabetes_y_test)
    predict = pred_model.predict(diabetes_X_test)
    print(predict)
    error=mean_squared_error(diabetes_X_test,predict)
    return error
    #raise NotImplementedError()
my_reg_model_mse(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test)   
```

Output:
```
[224.70261326  92.9629219  149.76810992  91.75430088  99.00602701
 143.72500481 236.78882348 100.21464803  73.62498556 102.63189008
 199.32157181  69.99912249 138.89052073 111.09223723  54.28704921
 159.4370781  119.55258438 119.55258438 181.19225649  55.49567023]
18579.38424160217
```

<!-- grade_id: my_reg_model_mse_test, points: 5, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_reg_model_mse(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), 
                    2548.07, rel_tol = 0.05)
```

Output:
```
[224.70261326  92.9629219  149.76810992  91.75430088  99.00602701
 143.72500481 236.78882348 100.21464803  73.62498556 102.63189008
 199.32157181  69.99912249 138.89052073 111.09223723  54.28704921
 159.4370781  119.55258438 119.55258438 181.19225649  55.49567023]
AssertionError                            Traceback (most recent call last)
<ipython-input-11-00c607b0c607> in <module>
      1 # Check you solution by running this cell
      2 import math
----> 3 assert math.isclose(my_reg_model_mse(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), 
      4                     2548.07, rel_tol = 0.05)

AssertionError:
```

---
## Question 2: 

#### Regression
Use the diabetes dataset from sklearn, which I have spilted into training and test sets (diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), please train the linear regression model to get the test accuracy and return `r2` (r sqaured value).  

<!-- grade_id: my_reg_model_r2_ans, points: 0, grade: false, solution: true -->
```python
def my_reg_model_r2(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test):
    """Write a function to train a machine learning model to get the `r2_score` on diabetes_X_test data . 
    Please use LinearRegression model to train a regressor. """
    # YOUR CODE HERE
    linear_regression = LinearRegression()
    linear_regression.fit(diabetes_X_train, diabetes_y_train)
    predictions = linear_regression.predict(diabetes_X_test)
    # Calculation of R2 Score
    return r2_score(diabetes_X_test,predictions)
my_reg_model_r2(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test)
    #raise NotImplementedError()
```

Output:
```
-10639807.760234658
```

<!-- grade_id: my_reg_model_r2_test, points: 4, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_reg_model_r2(diabetes_X_train, diabetes_y_train,diabetes_X_test,
                                    diabetes_y_test), 0.472575, rel_tol = 0.05)
```

Output:
```
/usr/local/lib/python3.9/dist-packages/sklearn/utils/validation.py in indexable(*iterables)
    441 
    442     result = [_make_indexable(X) for X in iterables]
--> 443     check_consistent_length(*result)
    444     return result
    445 

/usr/local/lib/python3.9/dist-packages/sklearn/utils/validation.py in check_consistent_length(*arrays)
    395     uniques = np.unique(lengths)
    396     if len(uniques) > 1:
--> 397         raise ValueError(
    398             "Found input variables with inconsistent numbers of samples: %r"
    399             % [int(l) for l in lengths]

ValueError: Found input variables with inconsistent numbers of samples: [422, 20]
```

```python

```